---
title       : Homework 2
description : CMSCI 120 Homework

--- type:NormalExercise lang:python xp:100 skills:1 key:b66b5626af
## First Programs

When learning any new programming language, one of the first programs is what is called "Hello World!".  We are going to use 1 of 6 primary building blocks for programs to create output to the screen.

Follow the instructions and complete instructions in the iPython shell to the right.

*** =instructions
- Complete the *print* statement with the phrase 'Hello World!'
- Press Submit Answer

*** =hint
- Enter Hello World! for the blanks in the print statement

*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
# Say Hello World!
print("____ ____")

```

*** =solution
```{python}
# Correct the syntax error below
print("Hello World!")
```

*** =sct
```{python}
# SCT written with pythonwhat: https://github.com/datacamp/pythonwhat/wiki

test_output_contains("Hello World!", pattern=False, no_output_msg = "Take a look at the instructions and try again")

success_msg("Nice job!!  Although a very simple program, that one line is a functional program.")

```

